# Introduction to Machine Learning: Deep Learning

**Instructor:** Daniel Acuna, Ph.D.
**Position:** Associate Professor of Computer Science
**Institution:** University of Colorado Boulder

---

Lab 2: Training and Regularizing Neural Networks

---

## Setup (do not edit)

```python
import pathlib
from typing import Tuple, Dict, List, Any

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.model_selection import train_test_split

# TensorFlow/Keras imports
import tensorflow as tf
from tensorflow.keras.models import Sequential
from tensorflow.keras.layers import Dense, Dropout, BatchNormalization
from tensorflow.keras.regularizers import l2
from tensorflow.keras.callbacks import EarlyStopping, ModelCheckpoint

# Suppress TensorFlow warnings for cleaner output
tf.get_logger().setLevel("ERROR")

RANDOM_STATE: int = 42
np.random.seed(RANDOM_STATE)
tf.random.set_seed(RANDOM_STATE)

# Fashion-MNIST class names for reference
CLASS_NAMES: List[str] = [
    "T-shirt/top", "Trouser", "Pullover", "Dress", "Coat",
    "Sandal", "Shirt", "Sneaker", "Bag", "Ankle boot"
]

_TRAIN_PATH = pathlib.Path("fashion_mnist_train.csv")
_TEST_PATH = pathlib.Path("fashion_mnist_test.csv")

if not _TRAIN_PATH.exists() or not _TEST_PATH.exists():
    raise FileNotFoundError(
        "Fashion-MNIST CSV files are missing from the lab directory. "
        "Please run the download script or ask the TA for assistance."
    )
```

## 1. Load and Explore the Dataset *(10 points)*

The Fashion-MNIST dataset contains 70,000 grayscale images of 10 clothing categories.
Each image is 28×28 pixels, flattened into 784 features.

Write a function `load_data()` that:
1. Reads the training CSV into a DataFrame
2. Returns the DataFrame

Store the shape of the training DataFrame in **`q1_shape`**.

<!-- grade_id: load-data, points: 0, grade: false, solution: true -->
```python


def load_data(path: pathlib.Path) -> pd.DataFrame:
    """Load a Fashion-MNIST CSV file.

    Parameters
    ----------
    path : pathlib.Path
        Path to the CSV file.

    Returns
    -------
    pd.DataFrame
        DataFrame containing pixel values and labels.
    """
    # your code here
    raise NotImplementedError


# Load training data and compute answer
train_df = load_data(_TRAIN_PATH)
q1_shape: Tuple[int, int] = train_df.shape
```

<!-- grade_id: load-data-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 1
assert isinstance(q1_shape, tuple), (
    "q1_shape must be a tuple. Use df.shape which returns (rows, cols)."
)
assert len(q1_shape) == 2, (
    "q1_shape should have 2 elements (rows, cols)."
)
assert q1_shape[0] > 0 and q1_shape[1] > 0, (
    "Shape values must be positive. Is your CSV loading correctly?"
)
print(f"Training dataset shape: {q1_shape}")
```

## 2. Preprocess the Data *(10 points)*

Neural networks work best with normalized inputs. Prepare the data by:

1. Separating features (pixel columns) from the target (label column)
2. Normalizing pixel values to the range [0, 1] by dividing by 255
3. Splitting into train (60%), validation (20%), and test (20%) sets

Store:
- **`X_train`**, **`X_val`**, **`X_test`**: Normalized feature arrays
- **`y_train`**, **`y_val`**, **`y_test`**: Label arrays
- **`q2_split_counts`**: Tuple of (n_train, n_val, n_test)

**Hint:** Use `train_test_split` twice with `random_state=RANDOM_STATE`.

<!-- grade_id: preprocess, points: 0, grade: false, solution: true -->
```python
# Load test data as well
test_df = load_data(_TEST_PATH)

# your code here
raise NotImplementedError
```

<!-- grade_id: preprocess-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 2
assert isinstance(q2_split_counts, tuple), (
    "q2_split_counts must be a tuple (n_train, n_val, n_test)."
)
assert len(q2_split_counts) == 3, "Tuple should have exactly 3 elements."
n_train, n_val, n_test = q2_split_counts
assert all(n > 0 for n in q2_split_counts), (
    "All splits must have positive sample counts."
)
assert n_train > n_val and n_train > n_test, (
    "Training set should be the largest. Check your split ratios."
)
assert X_train.max() <= 1.0 and X_train.min() >= 0.0, (
    "Pixel values should be normalized to [0, 1]. Did you divide by 255?"
)
print(f"Train/Val/Test split: {n_train} / {n_val} / {n_test}")
```

## 3. Build a Baseline Model (No Regularization) *(10 points)*

Build a neural network that is prone to overfitting (we'll fix it later!).

Architecture:
- **Input:** 784 features
- **Hidden Layer 1:** 256 neurons, ReLU activation
- **Hidden Layer 2:** 128 neurons, ReLU activation
- **Hidden Layer 3:** 64 neurons, ReLU activation
- **Output Layer:** 10 neurons, Softmax activation (for 10 classes)

Compile with:
- Optimizer: `'adam'`
- Loss: `'sparse_categorical_crossentropy'`
- Metrics: `['accuracy']`

Store the model in **`model_baseline`** and the parameter count in **`q3_n_params`**.

<!-- grade_id: baseline-model, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: baseline-model-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 3
assert isinstance(model_baseline, Sequential), (
    "model_baseline should be a Keras Sequential model."
)
assert isinstance(q3_n_params, (int, np.integer)), (
    "q3_n_params must be an integer. Use model.count_params()."
)
assert len(model_baseline.layers) == 4, (
    "Model should have 4 layers (3 hidden + 1 output)."
)
assert q3_n_params > 200000, (
    "This large model should have many parameters. Check your layer sizes."
)
print(f"Baseline model parameters: {q3_n_params:,}")
```

## 4. Train the Baseline Model *(10 points)*

Train the baseline model for **30 epochs** with `batch_size=128`.
Use the validation data to monitor performance.

Store:
- **`history_baseline`**: The training history object
- **`q4_final_val_acc`**: Final validation accuracy (last epoch), rounded to 3 decimals

After training, plot the training and validation accuracy curves.

<!-- grade_id: train-baseline, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: train-baseline-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 4
assert hasattr(history_baseline, 'history'), (
    "history_baseline should be a Keras History object. Did you call model.fit()?"
)
assert 'accuracy' in history_baseline.history, (
    "History should contain 'accuracy'. Did you include metrics=['accuracy']?"
)
assert 'val_accuracy' in history_baseline.history, (
    "History should contain 'val_accuracy'. Did you pass validation_data?"
)
assert isinstance(q4_final_val_acc, float), (
    "q4_final_val_acc must be a float."
)
assert 0 <= q4_final_val_acc <= 1, (
    "Accuracy must be between 0 and 1."
)
print(f"Final validation accuracy (baseline): {q4_final_val_acc:.3f}")
```

```python
# Visualize baseline training curves
plt.figure(figsize=(10, 4))
plt.subplot(1, 2, 1)
plt.plot(history_baseline.history['accuracy'], label='Train')
plt.plot(history_baseline.history['val_accuracy'], label='Validation')
plt.xlabel('Epoch')
plt.ylabel('Accuracy')
plt.title('Baseline Model Accuracy')
plt.legend()

plt.subplot(1, 2, 2)
plt.plot(history_baseline.history['loss'], label='Train')
plt.plot(history_baseline.history['val_loss'], label='Validation')
plt.xlabel('Epoch')
plt.ylabel('Loss')
plt.title('Baseline Model Loss')
plt.legend()
plt.tight_layout()
plt.show()
```

## 5. Add L2 Regularization *(10 points)*

L2 regularization (weight decay) penalizes large weights, encouraging simpler models.

Build a new model with the **same architecture** as the baseline, but add
`kernel_regularizer=l2(0.001)` to each Dense layer (except the output layer).

Store:
- **`model_l2`**: The L2-regularized model
- **`q5_l2_val_acc`**: Final validation accuracy after 30 epochs, rounded to 3 decimals

<!-- grade_id: l2-model, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: l2-model-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 5
assert isinstance(model_l2, Sequential), (
    "model_l2 should be a Keras Sequential model."
)
assert isinstance(q5_l2_val_acc, float), (
    "q5_l2_val_acc must be a float."
)
assert 0 <= q5_l2_val_acc <= 1, (
    "Accuracy must be between 0 and 1."
)
# Check that L2 regularization is applied
has_l2 = any(
    layer.kernel_regularizer is not None 
    for layer in model_l2.layers 
    if hasattr(layer, 'kernel_regularizer')
)
assert has_l2, (
    "Model should have L2 regularization. Use kernel_regularizer=l2(0.001)."
)
print(f"L2 Regularized model validation accuracy: {q5_l2_val_acc:.3f}")
```

## 6. Add Dropout *(10 points)*

Dropout randomly sets a fraction of activations to zero during training,
preventing co-adaptation of neurons.

Build a new model with the **same architecture**, but add `Dropout(0.3)` layers
after each hidden Dense layer (not after the output layer).

Store:
- **`model_dropout`**: The dropout model
- **`q6_dropout_val_acc`**: Final validation accuracy after 30 epochs, rounded to 3 decimals

<!-- grade_id: dropout-model, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: dropout-model-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 6
assert isinstance(model_dropout, Sequential), (
    "model_dropout should be a Keras Sequential model."
)
assert isinstance(q6_dropout_val_acc, float), (
    "q6_dropout_val_acc must be a float."
)
assert 0 <= q6_dropout_val_acc <= 1, (
    "Accuracy must be between 0 and 1."
)
# Check that Dropout layers are present
has_dropout = any(isinstance(layer, Dropout) for layer in model_dropout.layers)
assert has_dropout, (
    "Model should have Dropout layers. Add Dropout(0.3) after Dense layers."
)
print(f"Dropout model validation accuracy: {q6_dropout_val_acc:.3f}")
```

## 7. Implement Early Stopping *(10 points)*

Early stopping monitors validation loss and stops training when it stops improving,
preventing overfitting by not training too long.

Build the **same baseline architecture** (no regularization) but train it with
an `EarlyStopping` callback:
- Monitor: `'val_loss'`
- Patience: `5` (wait 5 epochs before stopping)
- `restore_best_weights=True`

Train for a maximum of **100 epochs** (early stopping will likely stop sooner).

Store:
- **`model_early_stop`**: The model trained with early stopping
- **`q7_stopped_epoch`**: The epoch at which training stopped (integer)
- **`q7_early_val_acc`**: Final validation accuracy, rounded to 3 decimals

<!-- grade_id: early-stopping, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: early-stopping-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 7
assert isinstance(model_early_stop, Sequential), (
    "model_early_stop should be a Keras Sequential model."
)
assert isinstance(q7_stopped_epoch, (int, np.integer)), (
    "q7_stopped_epoch must be an integer."
)
assert isinstance(q7_early_val_acc, float), (
    "q7_early_val_acc must be a float."
)
assert q7_stopped_epoch < 100, (
    "Early stopping should have stopped before 100 epochs. "
    "Check your EarlyStopping callback configuration."
)
assert q7_stopped_epoch > 5, (
    "Training stopped too early. Make sure patience=5."
)
print(f"Training stopped at epoch: {q7_stopped_epoch}")
print(f"Early stopping validation accuracy: {q7_early_val_acc:.3f}")
```

## 8. Add Batch Normalization *(10 points)*

Batch Normalization normalizes layer inputs, which stabilizes training and can
act as mild regularization.

Build a new model with BatchNormalization after each Dense layer (before activation):
- Dense(256) → BatchNormalization → ReLU (use `activation=None` in Dense, then separate layer)
- Or simpler: Dense(256, activation='relu') → BatchNormalization (after activation)

Use the simpler approach: add `BatchNormalization()` after each hidden Dense+ReLU layer.

Store:
- **`model_batchnorm`**: The batch-normalized model
- **`q8_batchnorm_val_acc`**: Final validation accuracy after 30 epochs, rounded to 3 decimals

<!-- grade_id: batchnorm-model, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: batchnorm-model-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 8
assert isinstance(model_batchnorm, Sequential), (
    "model_batchnorm should be a Keras Sequential model."
)
assert isinstance(q8_batchnorm_val_acc, float), (
    "q8_batchnorm_val_acc must be a float."
)
assert 0 <= q8_batchnorm_val_acc <= 1, (
    "Accuracy must be between 0 and 1."
)
# Check that BatchNormalization layers are present
has_batchnorm = any(isinstance(layer, BatchNormalization) for layer in model_batchnorm.layers)
assert has_batchnorm, (
    "Model should have BatchNormalization layers."
)
print(f"Batch Normalization model validation accuracy: {q8_batchnorm_val_acc:.3f}")
```

## 9. Compare All Models on Test Set *(10 points)*

Evaluate all five models on the **test set** and compare their performance.

Store a dictionary **`q9_test_accuracies`** with keys:
- `'baseline'`: Test accuracy of baseline model
- `'l2'`: Test accuracy of L2 regularized model
- `'dropout'`: Test accuracy of dropout model
- `'early_stopping'`: Test accuracy of early stopping model
- `'batchnorm'`: Test accuracy of batch normalization model

All values should be floats rounded to 3 decimals.

Also store **`q9_best_model`**: The key of the model with highest test accuracy.

<!-- grade_id: compare-models, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: compare-models-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 9
assert isinstance(q9_test_accuracies, dict), (
    "q9_test_accuracies must be a dictionary."
)
expected_keys = {'baseline', 'l2', 'dropout', 'early_stopping', 'batchnorm'}
assert set(q9_test_accuracies.keys()) == expected_keys, (
    f"Dictionary should have keys: {expected_keys}. Got: {set(q9_test_accuracies.keys())}"
)
for key, val in q9_test_accuracies.items():
    assert isinstance(val, float), f"Value for '{key}' must be a float."
    assert 0 <= val <= 1, f"Accuracy for '{key}' must be between 0 and 1."
assert isinstance(q9_best_model, str), (
    "q9_best_model must be a string (the model name)."
)
assert q9_best_model in expected_keys, (
    f"q9_best_model must be one of {expected_keys}."
)
print("Test Accuracies:")
for name, acc in q9_test_accuracies.items():
    print(f"  {name}: {acc:.3f}")
print(f"\nBest model: {q9_best_model}")
```

## 10. Analyze Overfitting Gap *(10 points)*

The **overfitting gap** is the difference between training accuracy and validation accuracy.
A large gap indicates the model is memorizing training data rather than learning general patterns.

For each model, compute the overfitting gap using the **final epoch's** accuracies:
`gap = train_accuracy - val_accuracy`

Store a dictionary **`q10_overfit_gaps`** with the same keys as above.
All values should be floats rounded to 3 decimals.

Also store **`q10_least_overfit`**: The key of the model with the **smallest** overfitting gap.

<!-- grade_id: overfit-gap, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: overfit-gap-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 10
assert isinstance(q10_overfit_gaps, dict), (
    "q10_overfit_gaps must be a dictionary."
)
expected_keys = {'baseline', 'l2', 'dropout', 'early_stopping', 'batchnorm'}
assert set(q10_overfit_gaps.keys()) == expected_keys, (
    f"Dictionary should have keys: {expected_keys}."
)
for key, val in q10_overfit_gaps.items():
    assert isinstance(val, float), f"Value for '{key}' must be a float."
assert isinstance(q10_least_overfit, str), (
    "q10_least_overfit must be a string (the model name)."
)
assert q10_least_overfit in expected_keys, (
    f"q10_least_overfit must be one of {expected_keys}."
)
print("Overfitting Gaps (Train Acc - Val Acc):")
for name, gap in q10_overfit_gaps.items():
    print(f"  {name}: {gap:.3f}")
print(f"\nLeast overfitting model: {q10_least_overfit}")
```

```python
# Visualization: Compare all models
fig, axes = plt.subplots(2, 3, figsize=(15, 8))

histories = {
    'Baseline': history_baseline,
    'L2': history_l2,
    'Dropout': history_dropout,
    'Early Stopping': history_early,
    'Batch Norm': history_batchnorm
}

for idx, (name, hist) in enumerate(histories.items()):
    ax = axes[idx // 3, idx % 3]
    ax.plot(hist.history['accuracy'], label='Train')
    ax.plot(hist.history['val_accuracy'], label='Val')
    ax.set_title(f'{name}\nGap: {hist.history["accuracy"][-1] - hist.history["val_accuracy"][-1]:.3f}')
    ax.set_xlabel('Epoch')
    ax.set_ylabel('Accuracy')
    ax.legend()
    ax.set_ylim([0.7, 1.0])

# Hide the 6th subplot
axes[1, 2].axis('off')

plt.tight_layout()
plt.suptitle('Comparison of Regularization Techniques', y=1.02)
plt.show()
```

## Next Steps

Congratulations on completing the assignment! Before submitting:

1. Make sure all your cells run without errors.
2. Ensure you've answered all parts of each question.
3. If any autograder tests fail, revisit your answers.
