# Introduction to Machine Learning: Deep Learning

**Instructor:** Daniel Acuna, Ph.D.
**Position:** Associate Professor of Computer Science
**Institution:** University of Colorado Boulder

---

# Lab 4: Sequence Models and Attention

In this lab, you will explore sequential data processing using Recurrent Neural Networks (RNNs), Long Short-Term Memory networks (LSTMs), and Transformers. You will work with the IMDB dataset to perform sentiment analysis.

---

## Setup (do not edit)

```python
import pathlib
import pickle
from typing import Tuple, Dict, List, Any

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.model_selection import train_test_split

import tensorflow as tf
from tensorflow.keras.preprocessing.text import Tokenizer
from tensorflow.keras.preprocessing.sequence import pad_sequences
from tensorflow.keras.models import Sequential, Model
from tensorflow.keras.layers import (
    Embedding, SimpleRNN, LSTM, Dense, Bidirectional, 
    Input, Dropout, LayerNormalization, MultiHeadAttention
)

# Suppress TensorFlow warnings
tf.get_logger().setLevel("ERROR")

RANDOM_STATE: int = 42
np.random.seed(RANDOM_STATE)
tf.random.set_seed(RANDOM_STATE)

_DATA_PATH = pathlib.Path("imdb_reviews.csv")

if not _DATA_PATH.exists():
    raise FileNotFoundError(
        "imdb_reviews.csv is missing. Please run the download script (w4_download_datasets.py) "
        "or ask the TA for assistance."
    )
```

## 1. Load and Inspect the Dataset *(10 points)*

Load the dataset from `imdb_reviews.csv`.

1. Read the CSV into a DataFrame.
2. Display the first 5 rows (ungraded).
3. Store the shape of the DataFrame in **`q1_shape`**.
4. Store the class distribution (counts of each label) in **`q1_dist`**.

<!-- grade_id: load-data, points: 0, grade: false, solution: true -->
```python
def load_data(path: pathlib.Path) -> pd.DataFrame:
    """Load the dataset from CSV.
    
    Parameters
    ----------
    path : pathlib.Path
        Path to the CSV file.
        
    Returns
    -------
    pd.DataFrame
        The loaded dataset.
    """
    # your code here
    #raise NotImplementedError
    return pd.read_csv(path)

df = load_data(_DATA_PATH)
q1_shape: Tuple[int, int] = df.shape
q1_dist: pd.Series = df['label'].value_counts().sort_index()

print(f"Shape: {q1_shape}")
print(f"Class distribution:\n{q1_dist}")
df.head()
```

Output:
```
Shape: (10000, 2)
Class distribution:
label
0    4912
1    5088
Name: count, dtype: int64
                                                text  label
0  ? never mind the serious logic gaps never mind...      0
1  ? overall this movie is dreadful and should ha...      0
2  ? this movie is about ? a gladiator who is bro...      0
3  ? drawing restraint 9 ? matthew barney br br h...      0
4  ? this story had a good plot to it about four ...      0
```

<!-- grade_id: load-data-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
assert isinstance(df, pd.DataFrame), "df should be a DataFrame"
assert isinstance(q1_shape, tuple), "q1_shape must be a tuple"
assert len(q1_shape) == 2, "q1_shape must have 2 dimensions"
assert "text" in df.columns and "label" in df.columns, "DataFrame must have 'text' and 'label' columns"
assert isinstance(q1_dist, pd.Series), "q1_dist must be a pandas Series"

```

## 2. Tokenization and Padding *(10 points)*

Neural networks cannot work directly with text strings. We need to:
1. Tokenize: Convert words to integers.
2. Pad: Ensure all sequences have the same length.

**Task:**
1. Initialize a `Tokenizer` with `num_words=10000` and `oov_token='<OOV>'`.
2. Fit the tokenizer on `df['text']`.
3. Convert texts to sequences of integers.
4. Pad the sequences to `maxlen=100` using `padding='post'` and `truncating='post'`.
5. Split the data into training (80%) and test (20%) sets using `train_test_split` with `random_state=RANDOM_STATE`.

Store:
- **`X_train`**, **`X_test`**: Padded sequences (numpy arrays)
- **`y_train`**, **`y_test`**: Labels (numpy arrays)
- **`vocab_size`**: The actual vocabulary size (should be 10000)

*Note:* We use a smaller vocabulary and sequence length to keep training fast.

<!-- grade_id: preprocessing, points: 0, grade: false, solution: true -->
```python
# Global parameters
VOCAB_SIZE = 10000
MAX_LEN = 100
EMBEDDING_DIM = 32

# your code here
#raise NotImplementedError
tokenizer = Tokenizer(num_words=VOCAB_SIZE, oov_token="<OOV>")
tokenizer.fit_on_texts(df['text'])
sequences = tokenizer.texts_to_sequences(df['text'])
padded_sequences = pad_sequences(sequences, maxlen=MAX_LEN, padding='post', truncating='post')
labels = df['label'].values
X_train, X_test, y_train, y_test = train_test_split(
    padded_sequences, labels, test_size=0.2, random_state=RANDOM_STATE
)
print(f"Training set shape: {X_train.shape}, {y_train.shape}")
print(f"Testing set shape: {X_test.shape}, {y_test.shape}")
```

Output:
```
Training set shape: (8000, 100), (8000,)
Testing set shape: (2000, 100), (2000,)
```

<!-- grade_id: preprocessing-test, points: 10, grade: true, solution: false -->
```python
assert isinstance(X_train, np.ndarray), "X_train must be a numpy array"
assert X_train.shape[1] == MAX_LEN, f"Sequence length must be {MAX_LEN}"
assert X_train.shape[0] + X_test.shape[0] == len(df), "Train + Test size should equal total samples"
assert len(tokenizer.word_index) > 0, "Tokenizer should be fitted"

```

## 3. Build a Simple RNN Model *(10 points)*

Recurrent Neural Networks (RNNs) process sequences by maintaining a hidden state.

Build a model with:
1. `Input` layer: shape=(MAX_LEN,)
2. `Embedding` layer: input_dim=VOCAB_SIZE, output_dim=EMBEDDING_DIM
3. `SimpleRNN` layer: 32 units
4. `Dense` layer: 1 unit, sigmoid activation (for binary classification)

Use `model.add()` to add each layer. Start with `Input(shape=(MAX_LEN,))` to define the input shape.

Compile with `adam` optimizer, `binary_crossentropy` loss, and `accuracy` metric.

Store the model in **`model_rnn`**.

<!-- grade_id: simple-rnn, points: 0, grade: false, solution: true -->
```python
# your code here
model_rnn = Sequential()
model_rnn.add(Input(shape=(MAX_LEN,)))
model_rnn.add(Embedding(input_dim=VOCAB_SIZE, output_dim=EMBEDDING_DIM))
model_rnn.add(SimpleRNN(32))
model_rnn.add(Dense(1, activation='sigmoid'))

model_rnn.compile(
    optimizer='adam',
    loss='binary_crossentropy',
    metrics=['accuracy']
 )

model_rnn.summary()
```

Output:
```
[1mModel: "sequential_1"[0m

┏━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━┳━━━━━━━━━━━━━━━━━━━━━━━━┳━━━━━━━━━━━━━━━┓
┃[1m [0m[1mLayer (type)                   [0m[1m [0m┃[1m [0m[1mOutput Shape          [0m[1m [0m┃[1m [0m[1m      Param #[0m[1m [0m┃
┡━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━╇━━━━━━━━━━━━━━━━━━━━━━━━╇━━━━━━━━━━━━━━━┩
│ embedding_1 ([38;5;33mEmbedding[0m)         │ ([38;5;45mNone[0m, [38;5;34m100[0m, [38;5;34m32[0m)        │       [38;5;34m320,000[0m │
├─────────────────────────────────┼────────────────────────┼───────────────┤
│ simple_rnn_2 ([38;5;33mSimpleRNN[0m)        │ ([38;5;45mNone[0m, [38;5;34m32[0m)             │         [38;5;34m2,080[0m │
├─────────────────────────────────┼────────────────────────┼───────────────┤
│ dense_1 ([38;5;33mDense[0m)                 │ ([38;5;45mNone[0m, [38;5;34m1[0m)              │            [38;5;34m33[0m │
└─────────────────────────────────┴────────────────────────┴───────────────┘

[1m Total params: [
```

<!-- grade_id: simple-rnn-test, points: 10, grade: true, solution: false -->
```python
assert isinstance(model_rnn, Sequential), "model_rnn must be a Sequential model"
assert len(model_rnn.layers) >= 3, "Model should have at least 3 layers"
assert isinstance(model_rnn.layers[0], Embedding), "First layer must be Embedding"
assert isinstance(model_rnn.layers[1], SimpleRNN), "Second layer must be SimpleRNN"
assert model_rnn.output_shape == (None, 1), "Output shape should be (None, 1)"

```

## 4. Train the Simple RNN *(10 points)*

Train `model_rnn` for **5 epochs** with `batch_size=64` using `X_train` and `y_train`.
Use `X_test` and `y_test` for validation.

Store the history in **`history_rnn`** and the final validation accuracy (rounded to 3 decimals) in **`q4_val_acc`**.

<!-- grade_id: train-rnn, points: 0, grade: false, solution: true -->
```python
# your code here
history_rnn = model_rnn.fit(
    X_train, y_train,
    epochs=5,
    batch_size=64,
    validation_data=(X_test, y_test),
    verbose=1
)

q4_val_acc = round(history_rnn.history['val_accuracy'][-1], 3)
```

Output:
```
Epoch 1/5
[1m 11/125[0m [32m━[0m[37m━━━━━━━━━━━━━━━━━━━[0m [1m1s[0m 10ms/step - accuracy: 0.9063 - loss: 0.2376
```

<!-- grade_id: train-rnn-test, points: 10, grade: true, solution: false -->
```python
assert 'val_accuracy' in history_rnn.history, "History should contain val_accuracy"
assert len(history_rnn.history['loss']) == 5, "Should train for 5 epochs"
print(f"SimpleRNN Val Acc: {q4_val_acc}")

```

## 5. Build an LSTM Model *(10 points)*

LSTMs are better at capturing long-term dependencies than SimpleRNNs.

Build a model with:
1. `Input` layer: shape=(MAX_LEN,)
2. `Embedding` layer: input_dim=VOCAB_SIZE, output_dim=EMBEDDING_DIM
3. `LSTM` layer: 32 units
4. `Dense` layer: 1 unit, sigmoid activation

Use `model.add()` to add each layer.

Compile with `adam`, `binary_crossentropy`, and `accuracy`.

Store the model in **`model_lstm`**.

<!-- grade_id: build-lstm, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
model_lstm.summary()
```

<!-- grade_id: build-lstm-test, points: 10, grade: true, solution: false -->
```python
assert isinstance(model_lstm, Sequential), "model_lstm must be a Sequential model"
assert isinstance(model_lstm.layers[1], LSTM), "Second layer must be LSTM"

```

## 6. Train the LSTM *(10 points)*

Train `model_lstm` for **5 epochs** with `batch_size=64`.
Use validation data.

Store history in **`history_lstm`** and final validation accuracy in **`q6_val_acc`**.

<!-- grade_id: train-lstm, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: train-lstm-test, points: 10, grade: true, solution: false -->
```python
assert 'val_accuracy' in history_lstm.history, "History should contain val_accuracy"
print(f"LSTM Val Acc: {q6_val_acc}")

```

## 7. Bidirectional LSTM *(10 points)*

Bidirectional RNNs process the sequence in both directions (forward and backward), often improving performance.

Build a model with:
1. `Input` layer: shape=(MAX_LEN,)
2. `Embedding` layer
3. `Bidirectional` wrapper around an `LSTM(32)` layer
4. `Dense` layer (sigmoid)

Use `model.add()` to add each layer.

Train for **5 epochs**.

Store model in **`model_bilstm`**, history in **`history_bilstm`**, and final validation accuracy in **`q7_val_acc`**.

<!-- grade_id: bilstm, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: bilstm-test, points: 10, grade: true, solution: false -->
```python
assert isinstance(model_bilstm.layers[1], Bidirectional), "Second layer must be Bidirectional"
# Check that it wraps an LSTM (works with both Keras 2 and 3)
wrapped_layer = getattr(model_bilstm.layers[1], 'forward_layer', None) or getattr(model_bilstm.layers[1], 'layer', None)
assert wrapped_layer is not None and isinstance(wrapped_layer, LSTM), "Wrapped layer must be LSTM"
print(f"Bidirectional LSTM Val Acc: {q7_val_acc}")

```

## 8. Transformer Block *(10 points)*

Transformers use self-attention to process sequences in parallel.

Create a custom `TransformerBlock` layer (or use Keras functional API) to build a small Transformer model.

Architecture:
1. Input layer (shape=`(MAX_LEN,)`)
2. Embedding layer
   - Use `Embedding(input_dim=VOCAB_SIZE, output_dim=EMBEDDING_DIM, input_length=MAX_LEN)`
3. MultiHeadAttention layer
   - Use `MultiHeadAttention(num_heads=2, key_dim=EMBEDDING_DIM)`
   - Note: In self-attention, the query, key, and value are the same tensor (the output of embedding).
     Pass `(embedding_output, embedding_output)` as arguments.
4. GlobalAveragePooling1D
   - To flatten the sequence into a single vector
5. Dropout (0.1)
6. Dense (20, relu)
7. Dropout (0.1)
8. Dense (1, sigmoid)

*Note: A full transformer usually includes LayerNormalization and skip connections, but we'll stick to a simplified version for this exercise.*

Implement this using the Functional API.

Store the model in **`model_transformer`**.

<!-- grade_id: transformer-build, points: 0, grade: false, solution: true -->
```python
from tensorflow.keras.layers import GlobalAveragePooling1D

# your code here
raise NotImplementedError
model_transformer.summary()
```

<!-- grade_id: transformer-build-test, points: 10, grade: true, solution: false -->
```python
assert isinstance(model_transformer, Model), "model_transformer must be a Keras Model"
# Check for MultiHeadAttention layer
assert any(isinstance(l, MultiHeadAttention) for l in model_transformer.layers), "Must include MultiHeadAttention layer"

```

## 9. Train Transformer *(10 points)*

Train the transformer model for **5 epochs** with `batch_size=64`.

Store history in **`history_transformer`** and final validation accuracy in **`q9_val_acc`**.

<!-- grade_id: train-transformer, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: train-transformer-test, points: 10, grade: true, solution: false -->
```python
print(f"Transformer Val Acc: {q9_val_acc}")
assert 'val_accuracy' in history_transformer.history

```

## 10. Model Comparison *(10 points)*

Compare the validation accuracies of all four models.

Create a dictionary **`q10_results`** with keys: `'SimpleRNN'`, `'LSTM'`, `'BiLSTM'`, `'Transformer'` and their corresponding final validation accuracies.

Which model performed best? Store the name of the best model in **`q10_best_model`**.

<!-- grade_id: comparison, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: comparison-test, points: 10, grade: true, solution: false -->
```python
assert isinstance(q10_results, dict), "q10_results must be a dictionary"
assert len(q10_results) == 4, "Must compare 4 models"
assert isinstance(q10_best_model, str), "q10_best_model must be a string"

print("Model Comparison:")
for model, acc in q10_results.items():
    print(f"{model}: {acc}")
print(f"Best Model: {q10_best_model}")
```

## Next Steps

Congratulations on completing the assignment! Before submitting:

1. Make sure all your cells run without errors.
2. Ensure you've answered all parts of each question.
3. If any autograder tests fail, revisit your answers.


```python

```