# Introduction to Machine Learning: Deep Learning

**Instructor:** Daniel Acuna, Ph.D.
**Position:** Associate Professor of Computer Science
**Institution:** University of Colorado Boulder

---

# Lab 5: Transformers and Vision Transformers

In this lab, you will explore the core components of the Transformer architecture, including
positional encoding, self-attention, and layer normalization. You will then build a simple
Vision Transformer (ViT) to classify Fashion MNIST images by treating image patches as tokens.

**Topics Covered:**
- Positional encoding
- Scaled dot-product attention
- Multi-head attention
- Layer normalization
- Vision Transformer (ViT) architecture
- Comparing encoder-only vs decoder-only Transformers

---

## Setup (do not edit)

```python
import pathlib
from typing import Tuple, Dict, List, Any

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.model_selection import train_test_split

import tensorflow as tf
from tensorflow.keras.models import Model
from tensorflow.keras.layers import (
    Input, Dense, Dropout, LayerNormalization, MultiHeadAttention,
    Layer, Embedding, Add, GlobalAveragePooling1D, Flatten, Reshape
)

# Suppress TensorFlow warnings
tf.get_logger().setLevel("ERROR")

RANDOM_STATE: int = 42
np.random.seed(RANDOM_STATE)
tf.random.set_seed(RANDOM_STATE)

_DATA_PATH = pathlib.Path("fashion_mnist_vit.csv")

if not _DATA_PATH.exists():
    raise FileNotFoundError(
        "fashion_mnist_vit.csv is missing. Please run the download script (w5_download_datasets.py) "
        "or ask the TA for assistance."
    )

# Class names for Fashion MNIST
CLASS_NAMES = ['T-shirt/top', 'Trouser', 'Pullover', 'Dress', 'Coat',
               'Sandal', 'Shirt', 'Sneaker', 'Bag', 'Ankle boot']
```

## 1. Load and Prepare Image Data *(10 points)*

Load the Fashion MNIST dataset and prepare it for the Vision Transformer.

**Tasks:**
1. Load the CSV file into a DataFrame.
2. Separate training and test data using the `split` column.
3. Extract labels (`y_train`, `y_test`) and pixel values.
4. Reshape images to 28×28 and normalize pixel values to [0, 1].
5. Store the training images in **`X_train`** and test images in **`X_test`**.
6. Store the shape of `X_train` in **`q1_shape`**.

<!-- grade_id: load-data, points: 0, grade: false, solution: true -->
```python
def load_and_prepare_data(path: pathlib.Path) -> Tuple[np.ndarray, np.ndarray, np.ndarray, np.ndarray]:
    """Load and prepare image data from CSV.
    
    Parameters
    ----------
    path : pathlib.Path
        Path to the CSV file.
        
    Returns
    -------
    Tuple[np.ndarray, np.ndarray, np.ndarray, np.ndarray]
        X_train, X_test, y_train, y_test
    """
    # your code here
    raise NotImplementedError

X_train, X_test, y_train, y_test = load_and_prepare_data(_DATA_PATH)
q1_shape: Tuple[int, int, int] = X_train.shape

print(f"Training data shape: {q1_shape}")
print(f"Test data shape: {X_test.shape}")
print(f"Number of classes: {len(np.unique(y_train))}")

# Visualize some samples
fig, axes = plt.subplots(1, 5, figsize=(12, 3))
for i, ax in enumerate(axes):
    ax.imshow(X_train[i], cmap='gray')
    ax.set_title(CLASS_NAMES[y_train[i]])
    ax.axis('off')
plt.suptitle("Sample Images from Fashion MNIST")
plt.tight_layout()
plt.show()
```

<!-- grade_id: load-data-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
assert isinstance(X_train, np.ndarray), "X_train must be a numpy array"
assert isinstance(q1_shape, tuple), "q1_shape must be a tuple"
assert len(q1_shape) == 3, "Images should be 3-dimensional (samples, height, width)"
assert X_train.max() <= 1.0 and X_train.min() >= 0.0, "Pixel values should be normalized to [0, 1]"
assert q1_shape[1] == 28 and q1_shape[2] == 28, "Images should be 28x28"

```

## 2. Positional Encoding *(10 points)*

Transformers process sequences in parallel and have no inherent notion of order.
Positional encoding adds position information to the input embeddings.

The original Transformer uses sinusoidal positional encoding:

$$PE_{(pos, 2i)} = \sin\left(\frac{pos}{10000^{2i/d_{model}}}\right)$$
$$PE_{(pos, 2i+1)} = \cos\left(\frac{pos}{10000^{2i/d_{model}}}\right)$$

**Task:**
Implement a function `create_positional_encoding(seq_len, d_model)` that returns
a positional encoding matrix of shape `(seq_len, d_model)`.

Store the positional encoding for `seq_len=49` and `d_model=64` in **`q2_pos_enc`**.

<!-- grade_id: positional-encoding, points: 0, grade: false, solution: true -->
```python
def create_positional_encoding(seq_len: int, d_model: int) -> np.ndarray:
    """Create sinusoidal positional encoding.
    
    Parameters
    ----------
    seq_len : int
        Sequence length (number of positions).
    d_model : int
        Embedding dimension.
        
    Returns
    -------
    np.ndarray
        Positional encoding of shape (seq_len, d_model).
    """
    # your code here
    raise NotImplementedError

# For ViT with 7x7 patches on 28x28 images: 49 patches
SEQ_LEN = 49
D_MODEL = 64

q2_pos_enc: np.ndarray = create_positional_encoding(SEQ_LEN, D_MODEL)

print(f"Positional encoding shape: {q2_pos_enc.shape}")

# Visualize the positional encoding
plt.figure(figsize=(10, 4))
plt.imshow(q2_pos_enc, aspect='auto', cmap='RdBu')
plt.xlabel('Embedding Dimension')
plt.ylabel('Position')
plt.title('Sinusoidal Positional Encoding')
plt.colorbar()
plt.show()
```

<!-- grade_id: positional-encoding-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
assert isinstance(q2_pos_enc, np.ndarray), "q2_pos_enc must be a numpy array"
assert q2_pos_enc.shape == (SEQ_LEN, D_MODEL), f"Shape should be ({SEQ_LEN}, {D_MODEL})"
assert np.abs(q2_pos_enc).max() <= 1.0, "Values should be between -1 and 1"

# Check that sin/cos pattern is correct
assert np.isclose(q2_pos_enc[0, 0], 0.0, atol=0.01), "PE[0,0] should be sin(0) = 0"
assert np.isclose(q2_pos_enc[0, 1], 1.0, atol=0.01), "PE[0,1] should be cos(0) = 1"

```

## 3. Scaled Dot-Product Attention *(10 points)*

The core of the Transformer is the attention mechanism:

$$\text{Attention}(Q, K, V) = \text{softmax}\left(\frac{QK^T}{\sqrt{d_k}}\right)V$$

**Task:**
Implement the `scaled_dot_product_attention` function.

Create test tensors and compute attention. Store:
- **`q3_attention_weights`**: The attention weights (after softmax)
- **`q3_output`**: The output of the attention mechanism

<!-- grade_id: attention, points: 0, grade: false, solution: true -->
```python
def scaled_dot_product_attention(
    query: np.ndarray, 
    key: np.ndarray, 
    value: np.ndarray
) -> Tuple[np.ndarray, np.ndarray]:
    """Compute scaled dot-product attention.
    
    Parameters
    ----------
    query : np.ndarray
        Query matrix of shape (seq_len, d_k).
    key : np.ndarray
        Key matrix of shape (seq_len, d_k).
    value : np.ndarray
        Value matrix of shape (seq_len, d_v).
        
    Returns
    -------
    Tuple[np.ndarray, np.ndarray]
        (output, attention_weights)
        - output: shape (seq_len, d_v)
        - attention_weights: shape (seq_len, seq_len)
    """
    # your code here
    raise NotImplementedError

# Test with simple example
np.random.seed(RANDOM_STATE)
test_seq_len = 4
test_d_k = 8

Q = np.random.randn(test_seq_len, test_d_k).astype(np.float32)
K = np.random.randn(test_seq_len, test_d_k).astype(np.float32)
V = np.random.randn(test_seq_len, test_d_k).astype(np.float32)

q3_output, q3_attention_weights = scaled_dot_product_attention(Q, K, V)

print(f"Attention weights shape: {q3_attention_weights.shape}")
print(f"Output shape: {q3_output.shape}")
print(f"Attention weights (rows sum to 1):\n{q3_attention_weights}")
```

<!-- grade_id: attention-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
assert q3_attention_weights.shape == (test_seq_len, test_seq_len), "Attention weights shape incorrect"
assert q3_output.shape == (test_seq_len, test_d_k), "Output shape incorrect"

# Check that attention weights sum to 1 along the key dimension
row_sums = np.sum(q3_attention_weights, axis=-1)
assert np.allclose(row_sums, 1.0, atol=1e-5), "Attention weights must sum to 1 for each query"

# Check that all weights are non-negative
assert np.all(q3_attention_weights >= 0), "Attention weights must be non-negative"

```

## 4. Layer Normalization *(10 points)*

Layer normalization normalizes across features for each sample:

$$\text{LayerNorm}(x) = \gamma \odot \frac{x - \mu}{\sqrt{\sigma^2 + \epsilon}} + \beta$$

Unlike batch normalization, layer norm:
- Normalizes across features (not batch)
- Works with any batch size
- Is ideal for sequence models

**Task:**
Implement `layer_norm(x, gamma, beta, epsilon)` that applies layer normalization.
Compute the mean and variance across the last axis (features).

Store the result of applying layer norm to a test array in **`q4_normalized`**.

<!-- grade_id: layer-norm, points: 0, grade: false, solution: true -->
```python
def layer_norm(
    x: np.ndarray, 
    gamma: np.ndarray, 
    beta: np.ndarray, 
    epsilon: float = 1e-6
) -> np.ndarray:
    """Apply layer normalization.
    
    Parameters
    ----------
    x : np.ndarray
        Input array of shape (..., features).
    gamma : np.ndarray
        Scale parameter of shape (features,).
    beta : np.ndarray
        Shift parameter of shape (features,).
    epsilon : float
        Small constant for numerical stability.
        
    Returns
    -------
    np.ndarray
        Normalized array, same shape as input.
    """
    # your code here
    raise NotImplementedError

# Test layer normalization
np.random.seed(RANDOM_STATE)
test_input = np.random.randn(3, 8).astype(np.float32) * 10 + 5  # Mean ~5, varied values
gamma = np.ones(8, dtype=np.float32)
beta = np.zeros(8, dtype=np.float32)

q4_normalized: np.ndarray = layer_norm(test_input, gamma, beta)

print(f"Input mean per sample: {test_input.mean(axis=-1)}")
print(f"Normalized mean per sample: {q4_normalized.mean(axis=-1)}")
print(f"Normalized std per sample: {q4_normalized.std(axis=-1)}")
```

<!-- grade_id: layer-norm-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
assert q4_normalized.shape == test_input.shape, "Shape should be preserved"

# After normalization with gamma=1, beta=0, each sample should have mean≈0, std≈1
sample_means = q4_normalized.mean(axis=-1)
sample_stds = q4_normalized.std(axis=-1)
assert np.allclose(sample_means, 0, atol=1e-5), "Normalized samples should have mean ≈ 0"
assert np.allclose(sample_stds, 1, atol=0.1), "Normalized samples should have std ≈ 1"

```

## 5. Patch Embedding for Vision Transformer *(10 points)*

Vision Transformers split images into fixed-size patches and linearly project each patch.

For a 28×28 image with 4×4 patches:
- Number of patches = (28/4) × (28/4) = 7 × 7 = 49
- Each patch has 4 × 4 = 16 pixels

**Task:**
Implement `extract_patches(images, patch_size)` that:
1. Splits each image into non-overlapping patches.
2. Flattens each patch into a vector.
3. Returns shape `(batch_size, num_patches, patch_dim)`.

Store the patches for the first 100 training images with `patch_size=4` in **`q5_patches`**.

<!-- grade_id: patch-embed, points: 0, grade: false, solution: true -->
```python
def extract_patches(images: np.ndarray, patch_size: int) -> np.ndarray:
    """Extract non-overlapping patches from images.
    
    Parameters
    ----------
    images : np.ndarray
        Images of shape (batch_size, height, width).
    patch_size : int
        Size of each square patch.
        
    Returns
    -------
    np.ndarray
        Patches of shape (batch_size, num_patches, patch_dim).
        patch_dim = patch_size * patch_size
    """
    # your code here
    raise NotImplementedError

PATCH_SIZE = 4
q5_patches: np.ndarray = extract_patches(X_train[:100], PATCH_SIZE)

print(f"Patches shape: {q5_patches.shape}")
print(f"Expected: (100, 49, 16)")

# Visualize patches from first image
fig, axes = plt.subplots(1, 7, figsize=(14, 2))
fig.suptitle("First 7 patches from first image (each 4x4)")
for i in range(7):
    patch = q5_patches[0, i].reshape(PATCH_SIZE, PATCH_SIZE)
    axes[i].imshow(patch, cmap='gray')
    axes[i].set_title(f"Patch {i}")
    axes[i].axis('off')
plt.tight_layout()
plt.show()
```

<!-- grade_id: patch-embed-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
expected_num_patches = (28 // PATCH_SIZE) ** 2  # 49
expected_patch_dim = PATCH_SIZE ** 2  # 16

assert q5_patches.shape == (100, expected_num_patches, expected_patch_dim), (
    f"Shape should be (100, {expected_num_patches}, {expected_patch_dim})"
)

# Check that patches reconstruct to original image
reconstructed = q5_patches[0].reshape(7, 7, 4, 4).transpose(0, 2, 1, 3).reshape(28, 28)
assert np.allclose(reconstructed, X_train[0]), "Patches should reconstruct to original image"

```

## 6. Build a Simple Vision Transformer *(15 points)*

Now let's build a simple Vision Transformer using Keras:

**Architecture:**
1. Input: Flattened patches of shape `(num_patches, patch_dim)`
2. Linear projection to `embed_dim=64`
3. Add positional encoding (as a Dense layer input)
4. Transformer encoder block:
   - `MultiHeadAttention` with `num_heads=4`, `key_dim=16`
   - Skip connection + `LayerNormalization`
   - `Dense(128, relu)` + `Dense(embed_dim)`
   - Skip connection + `LayerNormalization`
5. `GlobalAveragePooling1D` to aggregate patches
6. `Dense(64, relu)` + `Dropout(0.3)`
7. Output: `Dense(10, softmax)` for 10 classes

Store the model in **`model_vit`**.

<!-- grade_id: build-vit, points: 0, grade: false, solution: true -->
```python
# Model parameters
NUM_PATCHES = 49  # 7x7 patches
PATCH_DIM = 16    # 4x4 flattened
EMBED_DIM = 64
NUM_HEADS = 4
KEY_DIM = 16
FF_DIM = 128
NUM_CLASSES = 10

# your code here
raise NotImplementedError

model_vit.summary()
```

<!-- grade_id: build-vit-test, points: 15, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 15 points
assert isinstance(model_vit, Model), "model_vit must be a Keras Model"

# Check for required layers
layer_types = [type(layer).__name__ for layer in model_vit.layers]
assert 'MultiHeadAttention' in layer_types, "Model must include MultiHeadAttention"
assert 'LayerNormalization' in layer_types, "Model must include LayerNormalization"
assert 'GlobalAveragePooling1D' in layer_types, "Model must include GlobalAveragePooling1D"

# Check input and output shapes
assert model_vit.input_shape == (None, NUM_PATCHES, PATCH_DIM), "Input shape incorrect"
assert model_vit.output_shape == (None, NUM_CLASSES), "Output should have 10 classes"

```

## 7. Train the Vision Transformer *(10 points)*

Prepare the data and train the ViT model.

**Tasks:**
1. Extract patches from training and test images.
2. Train for **3 epochs** with `batch_size=64`.
3. Store the training history in **`history_vit`**.
4. Store the final validation accuracy (rounded to 3 decimals) in **`q7_val_acc`**.

<!-- grade_id: train-vit, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError

print(f"Final validation accuracy: {q7_val_acc}")
```

<!-- grade_id: train-vit-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
assert 'val_accuracy' in history_vit.history, "History should contain val_accuracy"
assert len(history_vit.history['loss']) == 3, "Should train for 3 epochs"
print(f"Vision Transformer Val Acc: {q7_val_acc}")

```

## 8. Compare with a Simple CNN *(10 points)*

Let's compare the ViT with a simple Convolutional Neural Network.

**Task:**
Build a simple CNN with:
1. `Input` layer: shape=(28, 28, 1)
2. `Conv2D(32, 3, relu, padding='same')` + `MaxPooling2D(2)`
3. `Conv2D(64, 3, relu, padding='same')` + `MaxPooling2D(2)`
4. `Flatten`
5. `Dense(64, relu)` + `Dropout(0.3)`
6. `Dense(10, softmax)`

Train for **3 epochs** with `batch_size=64`.

Store the model in **`model_cnn`**, history in **`history_cnn`**, and validation accuracy in **`q8_val_acc`**.

<!-- grade_id: build-cnn, points: 0, grade: false, solution: true -->
```python
from tensorflow.keras.layers import Conv2D, MaxPooling2D

# your code here
raise NotImplementedError

print(f"CNN Validation Accuracy: {q8_val_acc}")
```

<!-- grade_id: build-cnn-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
assert isinstance(model_cnn, Model), "model_cnn must be a Keras Model"
layer_types_cnn = [type(layer).__name__ for layer in model_cnn.layers]
assert 'Conv2D' in layer_types_cnn, "CNN must include Conv2D layers"
assert 'MaxPooling2D' in layer_types_cnn, "CNN must include MaxPooling2D layers"

print(f"CNN Val Acc: {q8_val_acc}")

```

## 9. Understanding BERT vs GPT Architecture *(5 points)*

Based on what you learned about BERT (encoder-only) and GPT (decoder-only):

**Task:**
Answer the following by setting the appropriate variables:

1. **`q9_bert_direction`**: BERT uses which type of attention?
   - `'bidirectional'` or `'causal'`

2. **`q9_gpt_direction`**: GPT uses which type of attention?
   - `'bidirectional'` or `'causal'`

3. **`q9_bert_task`**: BERT is better suited for which type of task?
   - `'understanding'` or `'generation'`

4. **`q9_gpt_task`**: GPT is better suited for which type of task?
   - `'understanding'` or `'generation'`

<!-- grade_id: bert-gpt, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError

print(f"BERT attention: {q9_bert_direction}")
print(f"GPT attention: {q9_gpt_direction}")
print(f"BERT best for: {q9_bert_task}")
print(f"GPT best for: {q9_gpt_task}")
```

<!-- grade_id: bert-gpt-test, points: 5, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 5 points
assert q9_bert_direction in ['bidirectional', 'causal'], "Must be 'bidirectional' or 'causal'"
assert q9_gpt_direction in ['bidirectional', 'causal'], "Must be 'bidirectional' or 'causal'"
assert q9_bert_task in ['understanding', 'generation'], "Must be 'understanding' or 'generation'"
assert q9_gpt_task in ['understanding', 'generation'], "Must be 'understanding' or 'generation'"

```

## 10. Model Comparison and Analysis *(10 points)*

Compare the performance and characteristics of the ViT and CNN models.

**Task:**
1. Create a dictionary **`q10_results`** with keys `'ViT'` and `'CNN'` containing their validation accuracies.
2. Determine which model performed better and store in **`q10_best_model`**.
3. Store the number of trainable parameters in each model:
   - **`q10_vit_params`**: Number of trainable parameters in ViT
   - **`q10_cnn_params`**: Number of trainable parameters in CNN

<!-- grade_id: comparison, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError

print("Model Comparison:")
print(f"  ViT: {q10_results['ViT']:.3f} accuracy, {q10_vit_params:,} parameters")
print(f"  CNN: {q10_results['CNN']:.3f} accuracy, {q10_cnn_params:,} parameters")
print(f"Best model: {q10_best_model}")

# Plot training curves
fig, axes = plt.subplots(1, 2, figsize=(12, 4))

axes[0].plot(history_vit.history['accuracy'], label='ViT Train')
axes[0].plot(history_vit.history['val_accuracy'], label='ViT Val')
axes[0].plot(history_cnn.history['accuracy'], label='CNN Train')
axes[0].plot(history_cnn.history['val_accuracy'], label='CNN Val')
axes[0].set_xlabel('Epoch')
axes[0].set_ylabel('Accuracy')
axes[0].set_title('Training Accuracy Comparison')
axes[0].legend()

axes[1].plot(history_vit.history['loss'], label='ViT Train')
axes[1].plot(history_vit.history['val_loss'], label='ViT Val')
axes[1].plot(history_cnn.history['loss'], label='CNN Train')
axes[1].plot(history_cnn.history['val_loss'], label='CNN Val')
axes[1].set_xlabel('Epoch')
axes[1].set_ylabel('Loss')
axes[1].set_title('Training Loss Comparison')
axes[1].legend()

plt.tight_layout()
plt.show()
```

<!-- grade_id: comparison-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
assert isinstance(q10_results, dict), "q10_results must be a dictionary"
assert 'ViT' in q10_results and 'CNN' in q10_results, "Must include both models"
assert isinstance(q10_best_model, str), "q10_best_model must be a string"
assert q10_best_model in ['ViT', 'CNN'], "Best model must be 'ViT' or 'CNN'"
assert isinstance(q10_vit_params, int), "Parameter count must be an integer"
assert isinstance(q10_cnn_params, int), "Parameter count must be an integer"

print("Model analysis complete!")
print(f"ViT parameters: {q10_vit_params:,}")
print(f"CNN parameters: {q10_cnn_params:,}")

```

## Next Steps

Congratulations on completing the assignment! Before submitting:

1. Make sure all your cells run without errors.
2. Ensure you've answered all parts of each question.
3. If any autograder tests fail, revisit your answers.