# Introduction to Machine Learning: Deep Learning

**Instructor:** Daniel Acuna, Ph.D.
**Position:** Associate Professor of Computer Science
**Institution:** University of Colorado Boulder

---

Lab 1: Neural Network Foundations and the Mathematics of Learning

---

## Setup (do not edit)

```python
import pathlib
from typing import Tuple, Dict, List, Any

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.model_selection import train_test_split
from sklearn.preprocessing import StandardScaler
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import accuracy_score

# TensorFlow/Keras imports
import tensorflow as tf
from tensorflow.keras.models import Sequential
from tensorflow.keras.layers import Dense

# Suppress TensorFlow warnings for cleaner output
tf.get_logger().setLevel("ERROR")

RANDOM_STATE: int = 42
np.random.seed(RANDOM_STATE)
tf.random.set_seed(RANDOM_STATE)

_DATA_PATH = pathlib.Path("breast_cancer_nn.csv")
if not _DATA_PATH.exists():
    raise FileNotFoundError(
        "breast_cancer_nn.csv is missing from the lab directory. Please download it or ask the TA "
        "for assistance."
    )
```

## 1. Load and Explore the Dataset *(10 points)*

Write a function `load_data()` that reads `breast_cancer_nn.csv` into a `pandas.DataFrame`.
Store the shape of the DataFrame in a variable called **`q1_shape`**.

The dataset contains 30 features computed from digitized images of breast mass fine needle
aspirates, plus a `target` column (0 = malignant, 1 = benign).

<!-- grade_id: load-data, points: 0, grade: false, solution: true -->
```python


def load_data() -> pd.DataFrame:
    """Load the Breast Cancer Wisconsin dataset from CSV.

    Returns
    -------
    pd.DataFrame
        A DataFrame containing the 30 features plus the target column.
    """
    # your code here
    raise NotImplementedError


# Compute the answer required by the autograder
q1_shape: Tuple[int, int] = load_data().shape
```

<!-- grade_id: load-data-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 1
assert isinstance(q1_shape, tuple), (
    "q1_shape must be a tuple. Use df.shape which returns (rows, cols)."
)
assert len(q1_shape) == 2, (
    "q1_shape should have 2 elements (rows, cols). "
    "Make sure you're using .shape, not .shape[0]."
)
assert q1_shape[0] > 0 and q1_shape[1] > 0, (
    "Shape values must be positive. Is your CSV loading correctly?"
)
print(f"Dataset shape: {q1_shape}")
```

## 2. Prepare Features and Target *(10 points)*

Separate the dataset into:
- **`X`**: A DataFrame containing all feature columns (everything except `target`)
- **`y`**: A Series containing the `target` column

Store the number of features in **`q2_n_features`** (an integer).

<!-- grade_id: prepare-data, points: 0, grade: false, solution: true -->
```python
df = load_data()

# your code here
raise NotImplementedError
```

<!-- grade_id: prepare-data-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 2
assert isinstance(q2_n_features, (int, np.integer)), (
    "q2_n_features must be an integer. Use X.shape[1] to get the number of columns."
)
assert q2_n_features > 0, (
    "Number of features must be positive. Did you create X correctly?"
)
assert "target" not in X.columns, (
    "The target column 'target' should not be in your feature matrix X. "
    "Use df.drop(columns='target') to remove it."
)
print(f"Number of features: {q2_n_features}")
```

## 3. Train/Validation/Test Split with Scaling *(10 points)*

Neural networks perform better when features are standardized (zero mean, unit variance).

1. Split the data into **training (60%)**, **validation (20%)**, and **test (20%)** sets.
   Use `random_state=RANDOM_STATE`.
2. Fit a `StandardScaler` on the **training data only**, then transform all three sets.

Store the row counts as a tuple **`q3_split_counts = (n_train, n_val, n_test)`**.

**Hint:** First split into train+val (80%) and test (20%), then split train+val into train (75%) and val (25%).

<!-- grade_id: split-scale, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: split-scale-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 3
assert isinstance(q3_split_counts, tuple), (
    "q3_split_counts must be a tuple like (n_train, n_val, n_test)."
)
assert len(q3_split_counts) == 3, "Tuple should have exactly 3 elements."
n_train, n_val, n_test = q3_split_counts
assert all(n > 0 for n in q3_split_counts), (
    "All splits must have positive row counts."
)
assert n_train > n_val and n_train > n_test, (
    "Training set should be the largest. Check your split ratios."
)
print(f"Train/Val/Test split: {n_train} / {n_val} / {n_test}")
```

## 4. Build a Neural Network Model *(10 points)*

Using Keras, build a neural network with the following architecture:
- **Input shape:** `(n_features,)` where `n_features` is from Question 2
- **Hidden Layer 1:** 16 neurons with ReLU activation
- **Hidden Layer 2:** 8 neurons with ReLU activation
- **Output Layer:** 1 neuron with Sigmoid activation (for binary classification)

Store the model in a variable called **`model`** and the total number of trainable
parameters in **`q4_n_params`** (an integer).

**Hint:** Use `model.count_params()` to get the parameter count.

<!-- grade_id: build-model, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: build-model-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 4
assert isinstance(model, Sequential), (
    "model should be a Keras Sequential model. "
    "Did you use Sequential([...])?"
)
assert isinstance(q4_n_params, (int, np.integer)), (
    "q4_n_params must be an integer. Use model.count_params()."
)
assert q4_n_params > 0, (
    "Model should have trainable parameters. Did you add layers?"
)
assert len(model.layers) == 3, (
    "Model should have exactly 3 layers (2 hidden + 1 output). "
    "Check your architecture."
)
print(f"Total trainable parameters: {q4_n_params}")
```

## 5. Compile and Train the Model *(10 points)*

1. **Compile** the model with:
   - Optimizer: `'adam'`
   - Loss: `'binary_crossentropy'`
   - Metrics: `['accuracy']`

2. **Train** the model for **50 epochs** with `batch_size=32` using the training data
   and validation data. Set `verbose=0` to suppress output.

Store the training history object in **`history`** and the final training accuracy
(last epoch) in **`q5_final_train_acc`** (float, rounded to 3 decimals).

<!-- grade_id: train-model, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: train-model-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 5
assert hasattr(history, 'history'), (
    "history should be a Keras History object. Did you call model.fit()?"
)
assert 'accuracy' in history.history, (
    "History should contain 'accuracy'. Did you include metrics=['accuracy'] in compile?"
)
assert 'val_accuracy' in history.history, (
    "History should contain 'val_accuracy'. Did you pass validation_data to fit()?"
)
assert isinstance(q5_final_train_acc, float), (
    "q5_final_train_acc must be a float."
)
assert 0 <= q5_final_train_acc <= 1, (
    "Accuracy must be between 0 and 1."
)
print(f"Final training accuracy: {q5_final_train_acc:.3f}")
```

## 6. Plot Training Curves *(10 points)*

Create a figure with **two subplots** side by side:
1. **Left subplot:** Training and validation **loss** over epochs
2. **Right subplot:** Training and validation **accuracy** over epochs

Both plots should include:
- X-axis label: "Epoch"
- Y-axis labels: "Loss" and "Accuracy" respectively
- Legend showing "Train" and "Validation"

Store the matplotlib figure object in **`q6_fig`**.

**Note:** Looking at these curves helps identify overfitting (when validation loss
increases while training loss continues to decrease).

<!-- grade_id: plot-curves, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: plot-curves-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 6
assert isinstance(q6_fig, plt.Figure), (
    "q6_fig must be a matplotlib Figure object. "
    "Use fig, axes = plt.subplots(...) and store fig in q6_fig."
)
assert len(q6_fig.axes) == 2, (
    "Figure should have exactly 2 subplots (loss and accuracy)."
)
print("Training curves plotted successfully!")
```

## 7. Evaluate on Test Set *(10 points)*

Evaluate the trained model on the **test set** (using the scaled test features).

Store the results as a tuple **`q7_test_metrics = (test_loss, test_accuracy)`**
with both values rounded to **3 decimals**.

<!-- grade_id: evaluate-test, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: evaluate-test-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 7
assert isinstance(q7_test_metrics, tuple), (
    "q7_test_metrics must be a tuple (loss, accuracy)."
)
assert len(q7_test_metrics) == 2, "Tuple should have exactly 2 elements."
test_loss, test_acc = q7_test_metrics
assert isinstance(test_loss, float) and isinstance(test_acc, float), (
    "Both loss and accuracy must be floats."
)
assert test_loss >= 0, "Loss must be non-negative."
assert 0 <= test_acc <= 1, "Accuracy must be between 0 and 1."
print(f"Test Loss: {test_loss:.3f}, Test Accuracy: {test_acc:.3f}")
```

## 8. Manual Forward Pass Function *(15 points)*

To reinforce your understanding of neural network mathematics, implement a function
that performs a **forward pass manually using NumPy**.

Implement **`manual_forward_pass(X: np.ndarray, weights: List) -> np.ndarray`** that:
1. Takes input data `X` (shape: n_samples × n_features)
2. Takes `weights` from `model.get_weights()` (a list of weight matrices and bias vectors)
3. Applies the forward pass through all layers:
   - For hidden layers: `Z = X @ W + b`, then `A = ReLU(Z)` where ReLU(x) = max(0, x)
   - For output layer: `Z = A @ W + b`, then `output = sigmoid(Z)`
4. Returns the final output probabilities

Store the predictions for the first 5 test samples in **`q8_manual_preds`** (shape: 5×1).

<!-- grade_id: manual-forward, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: manual-forward-test, points: 15, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 15 points
# Test Cell: Question 8
assert callable(manual_forward_pass), (
    "manual_forward_pass should be a callable function."
)
assert isinstance(q8_manual_preds, np.ndarray), (
    "q8_manual_preds must be a numpy array."
)
assert q8_manual_preds.shape == (5, 1), (
    f"q8_manual_preds should have shape (5, 1), got {q8_manual_preds.shape}. "
    "Make sure you're computing predictions for 5 samples with 1 output."
)
assert np.all((q8_manual_preds >= 0) & (q8_manual_preds <= 1)), (
    "Output probabilities must be between 0 and 1. Did you apply sigmoid to the output?"
)
print(f"Manual forward pass predictions (first 5):\n{q8_manual_preds.flatten()}")
```

## 9. Compare with Logistic Regression Baseline *(15 points)*

Train a `LogisticRegression` model (with `max_iter=1000` and `random_state=RANDOM_STATE`)
on the same scaled training data and evaluate it on the test set.

Store:
- **`q9_logreg_acc`**: Test accuracy of logistic regression (float, rounded to 3 decimals)
- **`q9_nn_better`**: Boolean indicating whether the neural network achieved higher
  test accuracy than logistic regression

**Note:** For simple datasets like this, logistic regression often performs comparably
to neural networks. Neural networks shine with more complex, high-dimensional data.

<!-- grade_id: compare-baseline, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: compare-baseline-test, points: 15, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 15 points
# Test Cell: Question 9
assert isinstance(q9_logreg_acc, float), (
    "q9_logreg_acc must be a float."
)
assert 0 <= q9_logreg_acc <= 1, (
    "Accuracy must be between 0 and 1."
)
assert isinstance(q9_nn_better, bool), (
    "q9_nn_better must be a boolean (True or False)."
)
print(f"Logistic Regression Test Accuracy: {q9_logreg_acc:.3f}")
print(f"Neural Network Test Accuracy: {q7_test_metrics[1]:.3f}")
print(f"Neural Network performed better: {q9_nn_better}")
```

## Next Steps

Congratulations on completing the assignment! Before submitting:

1. Make sure all your cells run without errors.
2. Ensure you've answered all parts of each question.
3. If any autograder tests fail, revisit your answers.


```python

```